# CS 3120/5120: Secure Distributed Computation
## Homework 2

```python
# Useful imports and utility functions
import pychor
import galois
GF = galois.GF(2**31-1)

@pychor.local_function
def share(secret, n):
    shares = [GF.Random() for _ in range(n-1)]
    shares.append(GF(secret) - GF(shares).sum())
    return shares
```

## Question 1 (10 points)

Implement an *outsourced summation protocol* with 2 servers and $n$ clients (where $n$ is large). In your protocol:
- Each of the $n$ clients has a secret input
- Each of the $n$ clients submits secret shares of their input to the 2 servers
- The 2 servers run a summation protocol to calculate the total sum

See the test case below for $n=10$ and $n=100$ examples of the use of the protocol.

<!-- grade_id: cell-3efcdde6842724c1, points: 0, grade: false, solution: true -->
```python
server1 = pychor.Party('server1')
server2 = pychor.Party('server2')

def protocol_outsourced_sum(inputs):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6d457cd296d09ce8, points: 10, grade: true, solution: false -->
```python
# TEST CASE 1
with pychor.LocalBackend():
    # Define the parties
    clients = [pychor.Party(f'client{i}') for i in range(10)]
    # Construct the input dict
    inputs = {p: p.constant(i) for i, p in enumerate(clients)}
    print('Inputs:', inputs)
    # Run the protocol
    result = protocol_outsourced_sum(inputs)
    print('Result:', result)
    # Check the results
    assert result.val == 45
    assert result.parties == set([server1, server2])
    
# TEST CASE 2
# Prints ommitted because they're long
with pychor.LocalBackend():
    clients = [pychor.Party(f'client{i}') for i in range(100)]
    inputs = {p: p.constant(i) for i, p in enumerate(clients)}
    result = protocol_outsourced_sum(inputs)
    print('Result:', result)
    assert result.val == 4950
    assert result.parties == set([server1, server2])
```

## Question 2 (10 points)

Write an ideal functionality for the outsourced computation protocol.

<!-- grade_id: cell-53578a32a7f3fba7, points: 0, grade: false, solution: true -->
```python
def functionality_outsourced_sum(inputs):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-30dbee54c8dfa540, points: 10, grade: true, solution: false -->
```python
# TEST CASE 1
with pychor.LocalBackend():
    # Define the parties
    clients = [pychor.Party(f'client{i}') for i in range(10)]
    # Construct the input dict
    inputs = {p: p.constant(i) for i, p in enumerate(clients)}
    print('Inputs:', inputs)
    # Run the protocol
    result = functionality_outsourced_sum(inputs)
    print('Result:', result)
    # Check the results
    assert result.val == 45
    assert result.parties >= set([server1, server2])
    
# TEST CASE 2
# Prints ommitted because they're long
with pychor.LocalBackend():
    clients = [pychor.Party(f'client{i}') for i in range(100)]
    inputs = {p: p.constant(i) for i, p in enumerate(clients)}
    result = functionality_outsourced_sum(inputs)
    print('Result:', result)
    assert result.val == 4950
    assert result.parties >= set([server1, server2])
```

## Question 3 (30 points)

Prove that `protocol_outsourced_sum` securely realizes `functionality_outsourced_sum` in the presence of a semi-honest adversary when `server1` is corrupted. Complete the proof by constructing a simulator via a hybrid argument. Reference [Chapter 3 of Programming MPC](https://jnear.github.io/programming-mpc/chapters/chapter03.html).

<!-- grade_id: cell-8185a60e99ac6bce, points: 30, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 4 (5 points)

How does the security of the outsourced computation model compare to the multiparty computation model (where the parties who provide the inputs also run the computation)? Consider both provable security and potential real-world incentives and harms.

<!-- grade_id: cell-cc8628ba7163f0e6, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE