# CS 3120/5120: Secure Distributed Computation
## Homework 1

```python
# Useful imports and utility functions
import pychor
import galois
GF = galois.GF(2**31-1)
```

## Question 1 (10 points)

Implement a 3-party summation protocol, `protocol_sum3`. The protocol should take 3 inputs, one per party, and add them together. It does not need to be secure.

<!-- grade_id: cell-ee91ecf45af4acb4, points: 0, grade: false, solution: true -->
```python
p1 = pychor.Party('p1')
p2 = pychor.Party('p2')
p3 = pychor.Party('p3')

def protocol_sum3(x, y, z):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3135ed9e0b939c73, points: 10, grade: true, solution: false -->
```python
# TEST CASE
with pychor.LocalBackend():
    x = p1.constant(5)
    y = p2.constant(6)
    z = p3.constant(7)
    
    result = protocol_sum3(x, y, z)
    print('Result:', result)
    assert result.val == 18
    assert result.parties == {p1, p2, p3}
```

## Question 2 (20 points)

Implement an $n$-party summation protocol. The inputs are given as a *dictionary* mapping parties to their input values. The dictionary may contain any number of items. Your protocol should return a single located value with their sum, located at all of the parties.

<!-- grade_id: cell-dbd3a9929036afc4, points: 0, grade: false, solution: true -->
```python
def protocol_sum_n(inputs):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a1a3ebda7b432568, points: 20, grade: true, solution: false -->
```python
# TEST CASE 1
with pychor.LocalBackend():
    # Define the parties
    parties = [pychor.Party(f'p{i}') for i in range(3)]
    # Construct the input dict
    inputs = {p: p.constant(i) for i, p in enumerate(parties)}
    print('Inputs:', inputs)
    # Run the protocol
    result = protocol_sum_n(inputs)
    print('Result:', result)
    # Check the results
    assert result.val == 3
    assert result.parties == set(parties)
    
# TEST CASE 2
# Prints ommitted because they're long
with pychor.LocalBackend():
    parties = [pychor.Party(f'p{i}') for i in range(30)]
    inputs = {p: p.constant(i) for i, p in enumerate(parties)}
    result = protocol_sum_n(inputs)
    assert result.val == 435
    assert result.parties == set(parties)
```

## Question 3 (20 points)

Implement an $n$-party *secure summation* protocol using additive secret sharing. The inputs are given as a *dictionary* mapping parties to their input values. The dictionary may contain any number of items. Your protocol should return a single located value with their sum, located at all of the parties.

<!-- grade_id: cell-696e8a276178baa9, points: 0, grade: false, solution: true -->
```python
def protocol_secure_sum_n(inputs):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-63b1ab43f091a757, points: 20, grade: true, solution: false -->
```python
# TEST CASE 1
with pychor.LocalBackend():
    parties = [pychor.Party(f'p{i}') for i in range(3)]
    inputs = {p: p.constant(i) for i, p in enumerate(parties)}
    print('Inputs:', inputs)
    result = protocol_secure_sum_n(inputs)
    print('Result:', result)
    assert result.val == GF(3)
    assert result.parties == set(parties)
    
# TEST CASE 2
# Prints ommitted because they're long
with pychor.LocalBackend():
    parties = [pychor.Party(f'p{i}') for i in range(30)]
    inputs = {p: p.constant(i) for i, p in enumerate(parties)}
    result = protocol_secure_sum_n(inputs)
    assert result.val == GF(435)
    assert result.parties == set(parties)
```

## Question 4 (5 points)

Why do we choose $p = 2^{31}-1$? Does it matter that $p$ is prime? What could happen if we choose $p$ too small, and the stream of numbers ends up being very large?

<!-- grade_id: cell-cc8628ba7163f0e6, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5 (5 points)

This system handles only unsigned (non-negative) integers. In 3-5 sentences, describe how we could extend it to handle *signed* integers?

<!-- grade_id: cell-21bb46fb153b3eca, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE