# CS 3120/5120: Secure Distributed Computation
## In-Class Exercise, week of 3/23/2026

## Useful Definitions for FV FHE Scheme

References: 
* [Somewhat Practical Fully Homomorphic Encryption](https://eprint.iacr.org/2012/144.pdf) (2012)
* [Homomorphic Encryption: a Toy Implementation in Python](https://bit-ml.github.io/blog/post/homomorphic-encryption-toy-implementation-in-python/)

```python
import numpy as np
from numpy.polynomial import polynomial as poly

# Constants for the system
n = 2**2
size = n
q = 2**14
modulus = q
t = 2
p = q**3
std1 = 1
std2 = 1

poly_mod = np.array([1] + [0] * (n - 1) + [1])

# Add and multiply polynomials, *without* modulus
def polyadd_wm(x, y, poly_mod):
    return poly.polydiv(poly.polyadd(x, y), poly_mod)[1]

def polymul_wm(x, y, poly_mod):
    return poly.polydiv(poly.polymul(x, y), poly_mod)[1]

# Add and multiply polynomials, *with* modulus
def polyadd(x, y, modulus, poly_mod):
    return np.int64(np.round(poly.polydiv(poly.polyadd(x, y) % modulus, poly_mod)[1] % modulus))

def polymul(x, y, modulus, poly_mod):
    return np.int64(np.round(poly.polydiv(poly.polymul(x, y) % modulus, poly_mod)[1] % modulus))

# Functions to generate important polynomials
def gen_binary_poly(size):
    return np.random.randint(0, 2, size, dtype=np.int64)

def gen_uniform_poly(size, modulus):
    return np.random.randint(0, modulus, size, dtype=np.int64)

def gen_normal_poly(size, mean, std):
    return np.int64(np.random.normal(mean, std, size=size))
```

## Question 1

Implement key generation for the somewhat homomorphic scheme in [FV12].

<!-- grade_id: cell-00a1a81e77eb0263, points: 0, grade: false, solution: true -->
```python
def keygen():
    # YOUR CODE HERE
    raise NotImplementedError()

keygen()
```

## Question 2

Implement encryption for the somewhat homomorphic scheme from [FV12].

<!-- grade_id: cell-ac906fcdfdad8278, points: 0, grade: false, solution: true -->
```python
def encrypt(pk, m):
    # YOUR CODE HERE
    raise NotImplementedError()
```

## Question 3

Implement decryption for the somewhat homomorphic scheme from [FV12].

<!-- grade_id: cell-9619226edb07375b, points: 0, grade: false, solution: true -->
```python
def decrypt(sk, ct):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-b03b8f6f96180d97, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
m_orig = np.array([1,0,0,0])
print('Message:', m_orig)
ct = encrypt(pk, m_orig)
print('Ciphertext:', ct)
m = decrypt(sk, ct)
print('Decrypted message:', m)
assert np.all(m == m_orig)
```

## Question 4

Implement addition for the somewhat homomorphic scheme of [FV12].

<!-- grade_id: cell-199cca6c84ee542e, points: 0, grade: false, solution: true -->
```python
def e_add(ct1, ct2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-02bce19dd026cd5c, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
m1 = np.array([0,0,1,1])
m2 = np.array([0,1,0,1])

ct1 = encrypt(pk, m1)
ct2 = encrypt(pk, m2)
ct3 = e_add(ct1, ct2)
m = decrypt(sk, ct3)
print('Addition result:', m)
assert np.all(polyadd(m1, m2, modulus, poly_mod) % t == m)
```

## Question 5

Implement multiplication for the somewhat homomorphic cryptosystem of FV12.

<!-- grade_id: cell-f7a1870c97744c01, points: 0, grade: false, solution: true -->
```python
def eval_keygen(sk):
    # YOUR CODE HERE
    raise NotImplementedError()

def e_mul(ct1, ct2, rlk):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-77a3a21341707596, points: 5, grade: true, solution: false -->
```python
# TEST CASE
sk, pk = keygen()
rlk = eval_keygen(sk)

m1 = np.array([0,0,1,1])
m2 = np.array([0,1,0,1])

ct1 = encrypt(pk, m1)
ct2 = encrypt(pk, m2)
ct4 = e_mul(ct1, ct2, rlk)
decrypted_m = decrypt(sk, ct4)
m1m2 = polymul(m1, m2, t, poly_mod)
print('Decrypted result:', decrypted_m)
print('Plaintext result:', m1m2)
assert np.all(decrypted_m == m1m2)
```