# CS 3120/5120: Secure Distributed Computation
## In-Class Exercise, Week of 1/12/2026

# Question 1

Write functions `share` and `reconstruct` for an additive secret sharing scheme over the integers. The functions should assume two secret shares.

<!-- grade_id: cell-16a7ad376675c884, points: 0, grade: false, solution: true -->
```python
import numpy as np

def share(x):
    # YOUR CODE HERE
    raise NotImplementedError()

def reconstruct(pair):
    s1, s2 = pair
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-03410747a69767fb, points: 5, grade: true, solution: false -->
```python
# Test case
assert reconstruct(share(5)) == 5
assert reconstruct(share(0)) == 0

a1, a2 = share(5)
assert reconstruct((a1 + 10, a2)) == 15

b1, b2 = share(20)
assert reconstruct((a1 + b1, a2 + b2)) == 25
```

# Question 2

Is the above secret sharing scheme *additively homomorphic*? Why or why not?

Is it *multiplicatively homomorphic*? Why or why not?

<!-- grade_id: cell-e28e0d5f448a5c30, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

# Question 3

What is problematic (for security) about the use of integers in the above solution?

<!-- grade_id: cell-52c5ebcaef37ac22, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4

Write functions `plusFE` and `multFE` that add and multiply two field elements in $GF(p)$, respectively.

<!-- grade_id: cell-c9adfe88526cc324, points: 0, grade: false, solution: true -->
```python
p = 97

def plusFE(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def multFE(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d1dfca2e7cebbcfb, points: 5, grade: true, solution: false -->
```python
# Test case
assert plusFE(2, 3) == 5
assert plusFE(50, 50) == 3
assert multFE(2, 5) == 10
assert multFE(6, 20) == 23
```

## Question 5

Write functions `share` and `reconstruct` for an additive secret sharing scheme over $GF(p)$. The functions should assume two secret shares.

<!-- grade_id: cell-7e41c12148039b16, points: 0, grade: false, solution: true -->
```python
def share(x):
    # YOUR CODE HERE
    raise NotImplementedError()

def reconstruct(shares):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-94eb8fd487f35605, points: 5, grade: true, solution: false -->
```python
# Test case
assert reconstruct(share(5)) == 5
assert reconstruct(share(0)) == 0

a1, a2 = share(5)
assert reconstruct((a1 + 10, a2)) == 15

b1, b2 = share(20)
assert reconstruct((a1 + b1, a2 + b2)) == 25
```

# Question 6

Write a protocol for two parties to sum their inputs. Each party provides a single input as an integer, and the parties send their numbers to each other so that both can calculate the result.

<!-- grade_id: cell-d4dab1c2f8d2697b, points: 0, grade: false, solution: true -->
```python
import pychor

p1 = pychor.Party('p1')
p2 = pychor.Party('p2')

def protocol_sum_int(v1, v2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-70812bff31305616, points: 5, grade: true, solution: false -->
```python
# Test case
with pychor.LocalBackend():
    # Define inputs for the protocol, located at the two parties
    v1 = p1.constant(5)
    v2 = p2.constant(3)
    
    # Run the protocol
    result = protocol_sum_int(v1, v2)

    # Check the results
    assert result.val == 8
    assert result.parties == {p1, p2}
```

# Question 7

Write a protocol for two parties to sum their inputs, but with secret sharing. Each party provides a single input as an integer. Each party sends *one secret share* of their number to the other. The parties sum the shares they hold, broadcast the results, and reconstruct the sum from the final shares.

<!-- grade_id: cell-cad48978b8f22848, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

def protocol_sum_secure(v1, v2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ed9564b597b351e5, points: 5, grade: true, solution: false -->
```python
# Test case
with pychor.LocalBackend():
    # Define inputs for the protocol, located at the two parties
    v1 = p1.constant(5)
    v2 = p2.constant(3)
    
    # Run the protocol
    result = protocol_sum_secure(v1, v2)

    # Check the results
    assert result.val == 8
    assert result.parties == {p1, p2}
```

# Question 8

Is `protocol_sum_secure` really secure? In what contexts? For what definition of "secure?"

<!-- grade_id: cell-bcaea15dc419bebf, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE