# CS3120/5120: Secure Distributed Computation
## In-Class Exercise, Week of 1/19/2026

```python
import pychor
import galois

GF = galois.GF((2**31)-1)
p1 = pychor.Party('p1')
p2 = pychor.Party('p2')
p3 = pychor.Party('p3')
```

## Question 1

Implement the functions `plusFE` and `multFE` using the `galois` library.

<!-- grade_id: cell-6b43d5c4cb3de8bf, points: 0, grade: false, solution: true -->
```python
def plusFE(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def multFE(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3f19513b4f19e7f0, points: 5, grade: true, solution: false -->
```python
a = GF(100)
b = GF(70)

assert plusFE(a, b) == GF(170)
assert plusFE(a, b) == 170
assert multFE(a, b) == GF(7000)
assert multFE(a, b) == 7000
```

## Question 2

Is the following protocol for 2-party summation secure? Why or why not?

```python
@pychor.local_function
def share(x):
    s1 = GF.Random()
    s2 = GF(x) - s1
    return s1, s2

def protocol_sum(x, y):
    x1, x2 = share(x).untup(2)
    y1, y2 = share(y).untup(2)
    x2.send(p1, p2)
    y1.send(p2, p1)
    sum1 = x1 + y1
    sum2 = x2 + y2
    sum1.send(p1, p2)
    sum2.send(p2, p1)
    total = sum1 + sum2
    return total

with pychor.LocalBackend():
    x = p1.constant(10)
    y = p2.constant(15)
    result = protocol_sum(x, y)
    print('Result:', result)
```

<!-- grade_id: cell-7867ffd80feb6347, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3

Extend the protocol to 3 parties.

<!-- grade_id: cell-f9734e3befc20b2c, points: 0, grade: false, solution: true -->
```python
@pychor.local_function
def share(x):
    # YOUR CODE HERE
    raise NotImplementedError()

def protocol_sum3(x, y, z):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-40714a04b0035104, points: 5, grade: true, solution: false -->
```python
with pychor.LocalBackend():
    x = p1.constant(10)
    y = p2.constant(15)
    z = p3.constant(23)
    result = protocol_sum3(x, y, z)
    print('Result:', result)
    assert result.val == GF(48)
```

## Question 4

Is the 3-party version secure? Why or why not?

<!-- grade_id: cell-5251195402cbd36d, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5

Write the ideal functionality for 2-party summation.

<!-- grade_id: cell-4db28d0fcc2aef42, points: 0, grade: false, solution: true -->
```python
def functionality_sum(x, y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-07e6b216153ce42e, points: 5, grade: true, solution: false -->
```python
with pychor.LocalBackend():
    x = p1.constant(10)
    y = p2.constant(15)
    result = functionality_sum(x, y)
    print('Result:', result)
```

## Question 6

Does `protocol_sum` securely realize `functionality_sum` in the presence of a semi-honest adversary, when P1 is corrupted? When P2 is corrupted?

<!-- grade_id: cell-6b354995ca42d459, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 7

Does `protocol_sum` securely realize `functionality_sum` in the presence of a *malicious* adversary? If not, why not?

<!-- grade_id: cell-07e156948f28d61b, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 8

Consider a semi-honest adversary. What information is included in P1's view of an execution of `protocol_sum`?

<!-- grade_id: cell-813b2fb7ee8bdac3, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 9

Describe the strategy for building a proof that `protocol_sum` securely realizes `functionality_sum` in the presence of a semi-honest adversary corrupting P1. Reference [Chapter 3 of Programming MPC](https://jnear.github.io/programming-mpc/chapters/chapter03.html).

<!-- grade_id: cell-368c746db8f210b1, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 10

Write a simulator proof that `protocol_sum` securely realizes `functionality_sum` in the presence of a semi-honest adversary corrupting P1.

<!-- grade_id: cell-8c2640f89e49076b, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 11

How would the proof differ when:
1. P2 is corrupt?
2. We use the 3-party protocol?
3. We use an $n$-party protocol?

<!-- grade_id: cell-5d17a11fa52ac908, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 12

If someone proves a protocol is secure, does that mean it's ok to use it?

<!-- grade_id: cell-ca9044ce231d3b4d, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE